Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Note that this Pre-class Work is estimated to take **34 minutes**.

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Norika Narimatsu"
COLLABORATORS = ""
```

---

# CS110 Pre-class Work - The Delta Median problem and Randomization

## Question 1 [time estimate: 16 minutes]

### Question 1a [time estimate: 7 minutes]

Write the worst sort function in the world, `worst_sort`. This function takes a list, and then randomly shuffles it until it happens to be in sorted order. Once it is in sorted order then the list is returned.


<!-- grade_id: cell-ead9c74054a1c5da, points: 0, grade: false, solution: true -->
```python
import random
def worst_sort(A):
    """
    Sort A in ascending order by randomly shuffling its 
    elements until they are in order.
    
    Parameters
    ----------
    A : list
        A list of numerical values
        
    Returns
    -------
    A : list
        Sorted list
        
    """
    # compare the input A and modified A, that is sorted by built in function 
    # while A is not equal to ascending ordered A,
    # continue to shuffle A
    while A != sorted(A):
        random.shuffle(A)
    return A
        
print(worst_sort([6,5,4,3,2,1,4,6,9,4,8,2]))
```

Output:
```
[1, 2, 2, 3, 4, 4, 4, 5, 6, 6, 8, 9]
```

<!-- grade_id: cell-ff4db3a3e8b04515, points: 1, grade: true, solution: false -->
```python
# Please ignore this cell. This cell is for us to implement the tests 
# to see if your code works properly. 
```

### Question 1b [time estimate: 3 minutes]
What is the best case complexity of this algorithm?

<!-- grade_id: cell-afac80171b3b6232, points: 0, grade: true, solution: true -->

I think best case complexity is O(n) when the input is orderdy sorted. With this input, we do not need to shuffle elements, but compare them to verigy their order is actually sorted.




### Question 1c [time estimate: 3 minutes]

What is the average case complexity?


<!-- grade_id: cell-9d52d34daa1e3478, points: 0, grade: true, solution: true -->
We have n! possible orders (permutations). Here, we shuffle A multiple times and its time complexityis O(n). Thus, average case complexity has exponential growth with O(n) * O(n!)


### Question 1d [time estimate: 3 minutes]

For what size lists is this a feasible method?

<!-- grade_id: cell-ab40f08768579798, points: 0, grade: true, solution: true -->
Since its time complexity exponentially increase, this method is feasible with short inputs. With small number of inputs, the number of permutation is also small too.


### [Optional] Question 1e [time estimate: 5 minutes]

Can you think of an even worse sorting method? In such a case, what would its complexity be? How big a list could you sort?

<!-- grade_id: cell-ac05806942341926, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

## Question 2 [time estimate: 18 minutes]

Approximate median finder. Given a list and $\tilde{\delta}$ (a number between 0 and 0.5), the approximate median finder function returns a number that is guaranteed to lie between the $(50-\tilde{\delta}/2)$% and $(50+\tilde{\delta}/2)$% percentiles, where $\tilde{\delta}=100\cdot \delta$. Implement such a function by randomly selecting an element from the list and testing whether or not it lies within the bounds. If it doesn’t then retry with a new random element, but only a limited number of retries are allowed (you decide the maximum number of retries.)

### Question 2a [time estimate: 5 minutes]
Complete the following function


<!-- grade_id: cell-2ab65512a6148d3c, points: 0, grade: false, solution: true -->
```python
def check_approx_med(A, num, delta):
    """
    Given a list and a number in a list, determine whether it is 
    within (50-delta/2)% and(50+delta/2)% percentiles.
    
    X% percentile of a list is defined to be the smallest number in 
    the list that is as large as at least X% numbers in the list. 
    
    Parameters
    ----------
    A : list
        A list of numerical values
    num 
        The number of interest
    delta : float
        The error, lies between 0 and 0.5

    Returns
    -------
    isApproxMed : bool
        A boolean value indicating whether num is whithin the bound
    
    """
    #decide the maximum number
    maximum = max(A)
    
    #where  delta hat =100* 𝛿 .
    delta_hat = delta * 100
    
    #setting percentile bounds as defined above
    L = ((50 - delta_hat/2)/100)* maximum
    U = (50 + delta_hat/2)/100 * maximum
    
    #  if num is between L and U, it returns true
    if L <= num & num<= U:
        return True
    return False

    #  for the test case 3, if A has one element, it will be true
    if A == [num]:
        return True
    
    return False

check_approx_med([0,5,3,2], 0, .25)
```

Output:
```
False
```

<!-- grade_id: cell-df56d7433abb6517, points: 1, grade: true, solution: false -->
```python
assert(check_approx_med([0, 1], 0.4, .25) == True)
assert(check_approx_med([0, 1], 0, .25) == False)
assert(check_approx_med([0], 0, .25) == True)
```

### Question 2b [time estimate: 5 minutes]
Complete the following function that makes use of `check_approx_med` above.


<!-- grade_id: cell-e19685b545147dde, points: 0, grade: false, solution: true -->
```python
def approx_med_finder(A, delta):
    """
    Given a list, find a number in the list that is between 
    50-delta/2% and 50+delta/2% percentiles
    
    Paramters
    ---------
    A : list
        A list of numerical values
    Delta : int
        The error, lies between 0 and 0.5
        
    Returns
    -------
    num
        The approximated median if it is found within 100//delta trials
        (this is one possible maximum number of retries. While you are 
        encouraged to play around with another limits, the submitted work 
        must use this number)
    None
        Finding failed, if nothing is found within 100//delta trials
   
   
    Note
    ----
    The 100//delta is chosen here as it is the expected number of trials for 
    the first successful finding to occur (using geometric distribution)
    
    """
    end = 100//delta
    trial = 0
    while trial < end:
       
        #select a random number from A
        nums = random.sample(A,1)
        num = nums[0]
        
        # if the num is within the bounds, it is median, so return that value
        if check_approx_med(A, num, delta) == True:
            return num
        
        #move on to the next trial
        trial+=1
        
    # if there is no value which falles between check_approx_med,
    #. return "Finding failed"
    return "Finding failed"
    
approx_med_finder([0,2],.25)
approx_med_finder([0,1,2],.25)

```

Output:
```
1
```

### Question 2c [time estimate: 3 minutes]

What is the probability of failure in each random trial? What is the probability of failure after all the allowed random trials? Does this scale with $\delta$ or $N$ (the number of elements in a list)?


<!-- grade_id: cell-efe7ab76f6a2546a, points: 0, grade: true, solution: true -->
complementary set
Success: the num is within the range
Failuer: the num is NOT within the range

P(failuer) = 1 - P(succes).
Depending on the value of delta (0 - 0.5), the bigger delta is, the wider the range of bounds will be -> more successful!

I don not think N (the number of elements in a list) influence the probability of success.






### Question 2d [time estimate: 5 minutes]
Analyze the expected runtime of `approx_med_finder`. Note that because the function uses `check_approx_med`, you will most likely need to analyze the runtime of that function, too.

<!-- grade_id: cell-86b199798477fc2a, points: 0, grade: true, solution: true -->

Best case, find median at the first try (the first element of the list is median). Its time complexity is O(1)
Worst case, check 100//delta times but there is no median. O(100//delta)