Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Note that this Pre-class Work is estimated to take **63 minutes**.

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Norika Narimatsu"
COLLABORATORS = ""
```

---

# CS110 Pre-class Work - Computational Applications of Randomized Algorithms

## Part A. Median Heap (watch this [video explanation](https://www.youtube.com/watch?v=756_8C9YBZQ&list=PLF_a-qBXTGFektoI6JUOTRL36JlvD04BR&index=5&t=0s) or read this [description](https://stackoverflow.com/a/15319593/7946759))

Throughout this pre-class work, please use the following definition of median: the median of a list of numbers is the one in the middle of the list when the list is ordered. When such the middle element can’t be determined (i.e., in a list of even length), the average of the two middle elements is the median. For example, 5 is the median of [-1,2,4,5,8,10,12], and (5+7)/2=6 is the median of [1,2,3,5,7,8,10,11].

Using the idea from Lesson Heaps and Heapsort, we can use a pair of heaps to create a data structure which allows fast access to the median. Use the heapq module in python to create both a max-heap and a min-heap. Note that by default, the heapq module in python only creates min-heaps, but if we multiply elements by -1 when we store them, then we can also create max-heaps.

## Question 1 [time estimate: 5 minutes]
Write a function `median(minh, maxh)`. It must return the median element.


<!-- grade_id: cell-43ce866dd1dc2d5b, points: 0, grade: false, solution: true -->
```python
# order
# max_heap | median | min_heap

def median(minh, maxh):
    '''''''''
    This function find the median by comparing 
    the length of max_heap and min_heap.
    
    Input:
    minh: min_heap, which is the right side of the median in the video 
    maxh: max_heap, which is the left side of the median in the video
    
    Output:
    median element
.
    '''''''''
    
    min_len = len(minh)
    max_len = len(maxh)
    diff = min_len - max_len

    # if their length is same
    if diff == 0:
        # calculate its average of two nodes
        return (maxh[0] + (-1*minh[0]))/2 
    elif diff > 0: # positive -> min is longer
        return -1*minh[0]
    else: # negative _> max is longer
        return maxh[0]


#testcase1 - 6
#maxh = [5,3,2,1]
#minh = [7,8,10,11]
#median(minh, maxh)

#testcase2 - 7
#maxh = [5,3,2,]
#minh = [7,8,10,11]
#median(minh, maxh)

#testcase3 -5
maxh = [5,3,2,1]
minh = [7,8,10]
median(minh, maxh)
```

Output:
```
5
```

## Question 2 [time estimate: 10 minutes]
Write a function `add_to_median_heap(minh, maxh, elem)`. It must accept a min heap, a max heap, and an element to add.


<!-- grade_id: cell-50cb08de70712ee7, points: 0, grade: false, solution: true -->
```python

import heapq


def add_to_median_heap(minh, maxh, elem):
    '''''''''
    This function adds new element to either min_heap or max_heap
    
    Input:
    minh: min_heap, which is the right side of the median in the video 
    maxh: max_heap, which is the left side of the median in the video
    elem: new element we will add 
    
    Output:
    None
    '''''''''
    min_len = len(minh)
    max_len = len(maxh)
    diff = min_len - max_len 
    
    # as we see in the video, if they are both empty, 
    # we can insert the element into 
    if min_len == 0 and max_len == 0: 
        heapq.heappush(minh,-1*elem)
        
    # if they have same length
    elif diff == 0:  
        # and also if the new element is smaller than or equal to the node of min_heap
        if elem <= -1*minh[0]:  
            # insert them into min_heap
            heapq.heappush(minh,-1*elem) 
        #if new element is bigger than the node of min_heap
        if elem > minh[0]: 
            # insert them into max_heap
            heapq.heappush(maxh,elem) 
    
    # keep balance
    #if minheap is longer than maxh
    elif diff > 0:
        # and also if the new element is smaller than or equal to the node of min_heap
        if elem <= -1*minh[0]:  
            # take the first element of min_heap and store as move
            move= -1*heapq.heappop(minh) 
            # puch this into max_heap to keep balance
            heapq.heappush(maxh,move) 
            # and then, puch the new element to balanced min_heap
            heapq.heappush(minh,-1*elem) 
            
        # if the new element is bigger than o the node of min_heap
        if elem > minh[0]: 
            # push it into maxh 
            heapq.heappush(maxh,elem) 
    #if minheap is shorter than maxh
    else:
        # and also if the new element is smaller than or equal to the node of min_heap
        if elem <= -1*minh[0]: 
            # put them into min_heap
            heapq.heappush(minh,-1*elem) 
        # if the new element is bigger than the node of min_heap
        if elem > -1*minh[0]:
            # take the first element of max_heap and store as move
            move=heapq.heappop(maxh)
            # move them into min_heap
            heapq.heappush(minh,-1*move)
            # and  push new element into maxheap
            heapq.heappush(maxh,elem)
```

<!-- grade_id: cell-673faa69bf330eda, points: 0, grade: true, solution: false -->
```python
# Please ignore this cell. This cell is for us to implement the tests 
# to see if your code works properly. 
```

## Question 3 [time estimate: 1 minute]

Uncomment and run the testing code given below to test your functions. It should print out numbers ranging from 1 to 50, in that order.

```python
minh = []
maxh = []
for a in range(1,100,2):
    add_to_median_heap(minh, maxh, a)
    print(median(minh, maxh))  
```

Output:
```
1
2.0
3
4.0
5
6.0
7
8.0
9
10.0
11
12.0
13
14.0
15
16.0
17
18.0
19
20.0
21
22.0
23
24.0
25
26.0
27
28.0
29
30.0
31
32.0
33
34.0
35
36.0
37
38.0
39
40.0
41
42.0
43
44.0
45
46.0
47
48.0
49
50.0
```

## Question 4 [time estimate: 3 minutes]
What’s the worst case complexity to build a median heap using `add_to_median_heap`?

<!-- grade_id: cell-c8e33ebdc8aaf328, points: 0, grade: true, solution: true -->
Worst case scenario: unbalance (skewed partition). In this casem we have to push new element into the heap which already has long length. For example, one heap has every elements, but the other has no elements.  We move the element to keep the balance untill the length will be n/2, which takes log(N) * N/2 times. To be precise, we have smaller notation to compare the length, O(1), but we can ignore it here. Thus, time complexity of add_to_median_heap is O( log(N) * N/2), which we can simplify as O(N log(N)).



## Question 5 [time estimate: 3 minutes] 
What’s the worst case complexity of `median`?

<!-- grade_id: cell-c061920faa337890, points: 0, grade: true, solution: true -->

I believe it is just O(1) because we need to compare the number and there is no loop inside our function..


## Question 6 [time estimate: 4 minutes]

How does this way of finding the median compare with the vanilla way of sorting the list and pick the middle element? Use arguments based on efficiency or clearity of the respective algorithms.

<!-- grade_id: cell-1f31afecda9730a7, points: 0, grade: true, solution: true -->
I think the code we write is doing much better job in terms of complexity. When we compare the worst case sceario, the sorting methods we leaned before takes O(N^2). But our function is (N log(N)). 

## [Optional] Question 7 [time estimate: 15 minutes]

Is it possible to extend this idea to any percentile? If it is, then write code to do so. If it’s not possible, prove why it is not possible.

<!-- grade_id: cell-4d0c197e6d484338, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

## Part B. Quickselect

Quicksort can be modified to find the $k$-th smallest element in an unordered list. This is known as quickselect. It does this by choosing a partition (as in quicksort). Once the list has been partitioned then we know how many elements lie to the left and to the right of the partition. This allows us to recursively call quickselect on the correct sublist.

## Question 1 [time estimate: 10 minutes]

Write a function `qselect(lst, k)`, which takes both a list and an index $k$. The function must then return the $k$-th smallest item in the list.


<!-- grade_id: cell-fb9c2bc9d5b4bc92, points: 0, grade: false, solution: true -->
```python
def qselect_withprint(lst,k):
    
    #set the last element of the list as a pivot
    pivot = lst[-1]  
    
    # reference https://note.nkmk.me/python-list-select-replace/
    # create two different list for left, which contains elements smaller than pivot, 
    # and right, which contains elements bigger than pivot
    left = [x for x in lst if x < pivot]
    print("left:", left)
    right = [y for y in lst if y > pivot]
    print("right:", right)

    
    if k < len(left):
        print("k < len(left)")
        # run function for left list 
        return qselect(left,k)
    elif k >= len(lst) - len(right):
        # run for right list
        print("k >= len(lst) - len(right)")
        return qselect(right, k - (len(lst)-len(right)))
    else:
        #keep numbers and results consistent
        print("else")
        return pivot
    
qselect([9,1,2,3,6,8,4,5], 5)
```

Output:
```
left: [1, 2, 3, 4]
right: [9, 6, 8]
k >= len(lst) - len(right)
left: [6]
right: [9]
k < len(left)
left: []
right: []
else
6
```

```python
def qselect(lst,k):
    
    #set the last element of the list as a pivot
    pivot = lst[-1]  
    
    # reference https://note.nkmk.me/python-list-select-replace/
    # create two different list for left, which contains elements smaller than pivot, 
    # and right, which contains elements bigger than pivot
    left = [x for x in lst if x < pivot]
    #print("left:", left)
    right = [y for y in lst if y > pivot]
    #print("right:", right)

    
    if k < len(left):
        #print("k < len(left)")
        # run function for left list 
        return qselect(left,k)
    elif k >= len(lst) - len(right):
        # run for right list
        #print("k >= len(lst) - len(right)")
        return qselect(right, k - (len(lst)-len(right)))
    else:
        #keep numbers and results consistent
        #print("else")
        return pivot
    
qselect([9,1,2,3,6,8,4,5], 5)
```

Output:
```
6
```

<!-- grade_id: cell-85a25ed6a1f62a66, points: 0, grade: true, solution: false -->
```python
import random
random.seed(123) # introducing a seed for reproducibility purposes
lst1 = list(range(100))
random.shuffle(lst1)
lst2 = []
for a in range(100):
    lst2.append(qselect(lst1, a))
assert(lst2 == list(range(100)))
```

## Question 2 [time estimate: 1 minute]

Uncomment and run the testing code given below to test your function. It should print out integers from 0 to 99.

```python

import random
random.seed(123) # introducing a seed for reproducibility purposes
lst = list(range(100))
random.shuffle(lst)
for a in range(100):
    print(qselect(lst, a))
```

Output:
```
0
1
2
3
4
5
6
7
8
9
10
11
12
13
14
15
16
17
18
19
20
21
22
23
24
25
26
27
28
29
30
31
32
33
34
35
36
37
38
39
40
41
42
43
44
45
46
47
48
49
50
51
52
53
54
55
56
57
58
59
60
61
62
63
64
65
66
67
68
69
70
71
72
73
74
75
76
77
78
79
80
81
82
83
84
85
86
87
88
89
90
91
92
93
94
95
96
97
98
99
```

## Question 3 [time estimate: 5 minutes]

Write down the recurrence relation for your code.

<!-- grade_id: cell-24144607b3ffade3, points: 0, grade: true, solution: true -->
From session 3.1: 
Whole problem = (How many we solve) T (n / (number of sub-problems)) + Theta( how long it takes us to set up the sub-problems + Combine )

T(N) = T(N/2) +O(N)

To get pertition, we compare the element -> O(N)
In the code, we have 2 cases (k> len(left) or len(right)) -> It means we  recursively call half case -> N/2.




## Question 4 [time estimate: 3 minutes]

Solve the recurrence relation for quickselect in the best case.


<!-- grade_id: cell-c125f4742933e820, points: 0, grade: true, solution: true -->
T(N) = O(1) to print the pivot and insert them.
There is one new element to insert.



## Question 5 [time estimate: 3 minutes]

Solve the recurrence relation for quickselect in the worst case.

<!-- grade_id: cell-852c3a71d8c2d7c1, points: 0, grade: true, solution: true -->
T(N ) = O(n^2) It happens when all elements are either larger or smaller than pivots. We have recurrence with one branch. 

