Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Note that this Pre-class Work is estimated to take **32 minutes**.

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Norika Narimatsu"
COLLABORATORS = ""
```

---

# CS110 Pre-class Work - Maximum subarray problem—part II

## Question 1 [time estimate: 2 minutes]
Paste in your Python implementation of the maximum subarray from the previous class in the cell below and use that to find out the value of the maximum subarray of this array: `A = [-2, -3, 4, -1, -2, 1]`

<!-- grade_id: cell-e69c8bbcd40ca242, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
def bruteforce_max_subarray(A):

    # YOUR CODE HERE
    max = 0 #the sum in the subarray
    index_i = 0 #starting point of the subarray
    index_j = 0 #ending point of the subarray

    for i in range(len(A)+1): 
        
        for j in range(i+1, len(A)+1): #since j will start from the next value of i -> i+1. i != j
            key = 0 #temporary value of new subarray to check later if it's greater than what we have
            
            for element in A[i: j]: #for elements in subarray, which starts from i and ends at j
                key = key + element #add all values in subarray

            if key > max: #check whether key value is larger than currently value.
                max = key  #If so, store value alongside all other information we're interested in
                index_i = i
                index_j = j
    #return desired values. 
    return index_i, index_j-1, max

print(bruteforce_max_subarray([-2,-3,4,-1,-2,1]))
print(bruteforce_max_subarray([-2,-3,4,-1,-2,1,8]))


```

Output:
```
(2, 2, 4)
(2, 6, 10)
```

```python

```

```python

```

```python

```

## Question 2 [time estimate: 5 minutes]
Now, your friend Joe comes and appends a single extra number at the end of the array, which becomes:  `B = [-2, -3, 4, -1, -2, 1, 8]`. With the output from ```max_subarray(A)```, do you need to rerun ```max_subarray()``` for the new array again to find the subarray we need? Explain why, or why not, in as much detail as possible, using, if necessary, a specific example to illustrate your reasoning. The subsequent questions will help you figure out an efficient algorithmic strategy to address the last question, but make sure to write your explanation above first, before answering the remaining questions.


<!-- grade_id: cell-d4b7cd845c816263, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE: I believe yes.
with the original array A,[-2,-3,4,-1,-2,1], the function gives us the output, (2,2,4), which means that the maximum subarray is just a single index [4]. Although there are 2 positive numbers, 4 and 1, in the original array, negative value between the two values decrease the maximum subarray. 

For this question, since we add 8 (positive number) at the end of the array, we cannot use the simple strategy as above (4 is nolonger the maximum subarray). To get this, we have to run brute force algorithm to check all the possible pairs of the starting point and the endpoint of the subarray.

## Question 3 [time estimate: 5 minutes]

**Determine if the following statement is True or False and explain your answer.**
If the maximum subarray of the array A is different than the maximum subarray of the array B (questions 1 and 2), the new maximum-subarray doesn’t need to include 8 (i.e., the newly appended element). 


<!-- grade_id: cell-9377964a8756f13b, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE: I am afraid that whether I understand this question properly. But I guess False, because we get the maximum subarray = 10 for the first time with 8. If we don't add 8, the maximum subarray is still 4.

## Question 4 [time estimate: 10 minutes]
Complete the Python function `incremental_max_subarray(x, mx)` in the cell below.

This [video](https://www.youtube.com/watch?v=AAgErqQmwmA&list=PLF_a-qBXTGFektoI6JUOTRL36JlvD04BR&index=4&t=0s)  might be helpful to understand the `incremental_max_subarray` problem.

<!-- grade_id: cell-0230e459f3d701f8, points: 0, grade: false, solution: true -->
```python
    """ 
    Parameters
    ----------
    x : list
        A NON-EMPTY list of numbers (e.g., the 
        array B in the first two questions above)
        
        If x has 1 element: returns the value of 
        the element regardless of the value of mx
        
    mx 
        The maximum subarray of x excluding its 
        ast element (i.e., compute the maximum 
        subarray of the input array x considering 
        only its first len(x) - 1 elements)
    
    Returns
    -------
    int
        The maximum subarray of the array x
    
    Example
    -------
    Using the array B in question 2, the result 
    of incremental_max_subarray(B, 4) is 10: 
    
        10 = 8 + 1 - 2 -1 + 4 

"""
def incremental_max_subarray(x, mx):
    #    If x has 1 element: returns the value of the element regardless of the value of mx, from pesudo
    if len(x) == 1:
        return x[0]y
    else:
        #initialize a value to hold the maximum value of the array
        current_sum = 0
        
        #set the low number for max array to store the first element of the list
        max_subarray = float("-inf")
        
        #iterate over every item in inputted list
        for i in reversed(range(-len(x),0)):

            #add the ith element to the current sum
            current_sum = current_sum +x[i]

            #if the count at this point is larger than previous maximum value, set count as new largest value
            if current_sum > max_subarray:
                max_subarray = current_sum
        
        #return the larger of maximum value without last entry, and maximum
        return max(max_subarray, mx)
    
B = [-2, -3, 4, -1, -2, 1, 8]
incremental_max_subarray(B, 2) 

    
```

Output:
```
10
```

```python

```

<!-- grade_id: cell-9abf19e56620ffa3, points: 1, grade: true, solution: false -->
```python
B = [-2, -3, 4, -1, -2, 1, 8]
assert(incremental_max_subarray(B, 4) == 10)
assert(incremental_max_subarray(B[:1], 0) == B[0])
```

```python

```

## Question 5

### Part 1 [time estimate: 5 minutes]

Now use `incremental_max_subarray(x, mx)` iteratively, starting from the first element, to calculate the maximum subarray of a given list. Come up with a few test cases to test your code.

<!-- grade_id: cell-fd96d4ccccd99fe6, points: 0, grade: false, solution: true -->
```python

    
    

def max_subarray(A):
    """
    Using `incremental_max_subarray` iteratively on A 
    to produce the value of the maximum subarray of A.
    
    Parameters
    ----------
    A : list
        A NON-EMPTY list of floats
        
    Returns
    -------
    float
        The sum of the maximum subarray of A 
    
    """
    #base-case
    #if len(A) == 1:
    global_sum = 0
    
    #running a for loop to add new values and restore the new global sum witht the function above, recursive 
    for a in range(len(A)+1):
        global_sum  = incremental_max_subarray(A, global_sum)    
    return global_sum

A = [-2, -3, 4, -1, -2, 1]
B = [-2, -3, 4, -1, -2, 1, 8]
assert(max_subarray(A) == 4)
assert(max_subarray(B) == 10)
```

Output:
```
AssertionError                            Traceback (most recent call last)
<ipython-input-7-04ff20845a7b> in <module>
     26 A = [-2, -3, 4, -1, -2, 1]
     27 B = [-2, -3, 4, -1, -2, 1, 8]
---> 28 assert(max_subarray(A) == 4)
     29 assert(max_subarray(B) == 10)

AssertionError:
```

```python

```

```python

```

```python

```

```python

```

```python

```

<!-- grade_id: cell-669ad779766aa2c3, points: 1, grade: true, solution: false -->
```python
# Please ignore this cell. This cell is for us to implement the tests 
# to see if your code works properly. 
```

### Part 2 [time estimate: 5 minutes]
Is this more efficient than the divide-and-conquer approach? Explain.

```python
Depends on the input?
Divide and conquer: O(nlogn)
Increment: O(n^2)
with the large input -> divide and conquer is better.


```

<!-- grade_id: cell-cd8f0b130a7136db, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE